---
title: "World Happiness Report"
output: 
  learnr::tutorial:
    progressive: true
    allow_skip: true
runtime: shiny_prerendered
description: >
  Learn how to explore the World Happiness Report dataset and make some pretty plots!
---

```{r setup, include=FALSE}
knitr::opts_chunk$set(echo = TRUE)
library(learnr)
gradethis::gradethis_setup()

library(tidyverse)
library(janitor)

happy_full <- read_csv("https://raw.githubusercontent.com/BAREJAA/website_for_john/master/datasets/world-happiness-report-2021.csv") %>% clean_names() %>% rename(region = regional_indicator)

happy_full %>% select(country_name:ladder_score, logged_gdp_per_capita:ladder_score_in_dystopia) %>% relocate(ladder_score_in_dystopia, .after = region) -> happy_select
```

## Analysing the World Happiness Report

In this exercise, you will apply what you've learned in class to perform exploratory data analysis (EDA) on the **World Happiness Report** and make some pretty plots.  

This dataset was downloaded from the website [Kaggle](https://www.kaggle.com/ajaypalsinghlo/world-happiness-report-2021?select=world-happiness-report.csv). We will use the **2021** data in this exercise. This dataset is stored in an object called `happy_full` 

###  

In this exercise, you will practice:

* gaining quick insight into the type of data this dataset contains
* using functions from the Dplyr package to wrangle your data and obtain useful summaries
* making pretty plots!

### Prerequisites

Please watch the recordings on Dplyr and Ggplot2  

## Take a look at your data

### What does the dataset look like?

A couple useful things to know about your dataset are -
- The number of rows and columns
- The types of variables the dataset contains  

What function can you use to get this information?

```{r glimpse, exercise = TRUE, exercise.eval = FALSE, exercise.cap = "Take a glimpse at your data"}

```

<div id="glimpse-hint">
**Hint:** Try the dplyr `glimpse()` function.
</div>

```{r glimpse-solution}
glimpse(happy_full)
```

```{r glimpse-code-check}
grade_code()
```

### Double data type

We can see that `happy_full` contains many variables that are of type double. 

```{r quiz-double, echo = FALSE}
quiz(
  question("The `double` data type refers to which of the following?",
    answer("A string"),
    answer("An integer"),
    answer("A number with a decimal point", correct = TRUE),
    answer("A factor"))
)
```

## Selecting Variables 

The `happy_full` dataset contains many variables. This gives us the chance to practice our `select()`-ing skills!

### Simple selects

Let's warm up by performing some basic select operations  

How would you select just the columns `region` and `ladder_score`?  

```{r simple-select, exercise = TRUE, exercise.eval = FALSE, exercise.cap = "A simple selection"}

```

```{r simple-select-solution}
happy_full %>% 
  select(region, ladder_score)
```

```{r simple-select-code-check}
grade_code()
```

Now select everything between (and including) `social_support` and `generosity`

```{r select-between, exercise = TRUE, exercise.eval = FALSE, exercise.cap = "Selecting multiple columns"}

```

```{r select-between-solution}
happy_full %>% 
  select(social_support:generosity)
```

```{r select-between-code-check}
grade_code()
```

### Slightly-more-difficult selects  

Let's try something more challenging now. Select all variables that **do not** have underscores in their names

```{r no-underscore, exercise = TRUE, exercise.eval = FALSE, exercise.cap = "No underscores"}

```

<div id="no-underscore-hint">
**Hint:** You'll need a helper function. Also, don't forget `!`
</div>

```{r no-underscore-solution}
happy_full %>% 
  select(!contains("_"))
```

```{r no-underscore-code-check}
grade_code()
```

Helper functions can be really...helpful! Ok, no we're ready to select the variables we will need for the rest of this exercise. Create a new dataframe called `happy_df` that contains the following variables (in the specified order!) - `country_name`, `region`, `ladder_score_in_dystopia`, `logged_gdp_per_capita`, `social_support`, `healthy_life_expectancy`, `freedom_to_make_life_choices`, `generosity`, `perceptions_of_corruption`  

Avoid simply typing out the names of all these variables. Add `glimpse(happy_df)` as the last line to see if you got the right answer.

```{r happy-select, exercise = TRUE, exercise.eval = FALSE, exercise.cap = "No underscores"}

```

<div id="happy-select-hint">
**Hint:** You'll need a helper function. Also, don't forget `!`
</div>

```{r happy-select-solution}
happy_df <- happy_full %>% 
                select(country_name:ladder_score, logged_gdp_per_capita:ladder_score_in_dystopia) %>%   
                relocate(ladder_score_in_dystopia, .after = region)

glimpse(happy_df)
```

```{r happy-select-code-check}
grade_code()
```

**Note:** `happy_df` hasn't actually been saved anywhere. So we will be using an identical dataset called `happy_select` for the rest of this exercise. 